## Problem 3: Reading coordinates from a file and creating geometries (*5 points*) 

One of the most typical problems in GIS is the situation where you have a set of coordinates in some file, and you need to map those. Python is a really handy tool for these kind of situations, as it is possible to read data from (basically) any kind of input datafile (such as csv-, txt-, excel-, gpx-files (gps data), databases etc.). 

Thus, let's see how we can read data from a file and create Point -objects from them that can be saved e.g. as a new Shapefile (we will learn this next week). Our dataset **[travelTimes_2015_Helsinki.txt](data/travelTimes_2015_Helsinki.txt)** consist of 
travel times between specific locations in Helsinki Region. The first four rows of our data looks like this:

```
   from_id;to_id;fromid_toid;route_number;at;from_x;from_y;to_x;to_y;total_route_time;route_time;route_distance
   5861326;5785640;5861326_5785640;1;08:10;24.9704379;60.3119173;24.8560344;60.399940599999994;125.0;99.0;22917.6
   5861326;5785641;5861326_5785641;1;08:10;24.9704379;60.3119173;24.8605682;60.4000135;123.0;102.0;23123.5
   5861326;5785642;5861326_5785642;1;08:10;24.9704379;60.3119173;24.865102;60.4000863;125.0;103.0;23241.3
```

As we can see, there exists many columns in the data, but the few important ones needed here are:

| Column | Description |
|--------|-------------|
| from_x | x-coordinate of the **origin** location (longitude) |
| from_y | y-coordinate of the **origin** location (latitude) |
| to_x   | x-coordinate of the **destination** location (longitude)|
| to_y   | y-coordinate of the **destination** location (latitude) |
| total_route_time | Travel time with public transportation at the route |

Read more about the input data set at the Digital Geography Lab / Accessibility Research Group (University of Helsinki, Finland) website: https://blogs.helsinki.fi/accessibility/helsinki-region-travel-time-matrix/.

### Steps

1: Read the [data/travelTimes_2015_Helsinki.txt](data/travelTimes_2015_Helsinki.txt) file into a variable **`data`** using  pandas.

**NOTE:** What is the separator in the data (see above)? Remember to take that into account when reading the data.


<!-- grade_id: read_data, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
import pandas as pd
from shapely.geometry import Point, LineString, Polygon

data = pd.read_csv('data/travelTimes_2015_Helsinki.txt', sep = ';')
# raise NotImplementedError()
```

```python
#Check how many rows and columns there are:
print(len(data))
print(len(data.columns))
```

Output:
```
14643
13
```

<!-- grade_id: problem_3_read_data_test, points: 0.5, grade: true, solution: false -->
```python
# NON-EDITABLE CODE CELL FOR TESTING YOUR SOLUTION

# This test print should print first five rows in the data (if not, something is incorrect):
print(data.head())
```

Output:
```
from_id    to_id      fromid_toid  route_number     at     from_x  \
0  5861326  5785640  5861326_5785640             1  08:10  24.970438   
1  5861326  5785641  5861326_5785641             1  08:10  24.970438   
2  5861326  5785642  5861326_5785642             1  08:10  24.970438   
3  5861326  5785643  5861326_5785643             1  08:10  24.970438   
4  5861326  5787544  5861326_5787544             1  08:10  24.970438   

      from_y       to_x       to_y  total_route_time  route_time  \
0  60.311917  24.856034  60.399941             125.0        99.0   
1  60.311917  24.860568  60.400014             123.0       102.0   
2  60.311917  24.865102  60.400086             125.0       103.0   
3  60.311917  24.869636  60.400159             129.0       107.0   
4  60.311917  24.842582  60.397478             118.0        92.0   

   route_distance  route_total_lines  
0         22917.6                2.0  
1         23123.5                2.0  
2         23241.3                2.0  
3    
```

2: Select the 4 columns that contain coordinate information (**'from_x'**, **'from_y'**, **'to_x'**, **'to_y'**) and store them in variable **`data`** (i.e. update the data -variable).


<!-- grade_id: select_data, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
data = data[['from_x','from_y','to_x', 'to_y']]
print(data)
# numbers seemes to be rounded by pandas!!
# raise NotImplementedError()
```

Output:
```
from_x     from_y       to_x       to_y
0      24.970438  60.311917  24.856034  60.399941
1      24.970438  60.311917  24.860568  60.400014
2      24.970438  60.311917  24.865102  60.400086
3      24.970438  60.311917  24.869636  60.400159
4      24.970438  60.311917  24.842582  60.397478
...          ...        ...        ...        ...
14638  24.970438  60.311917  24.559702  60.174754
14639  24.970438  60.311917  24.564204  60.174837
14640  24.970438  60.311917  24.555367  60.172428
14641  24.970438  60.311917  24.559868  60.172511
14642  24.970438  60.311917  24.555534  60.170185

[14643 rows x 4 columns]
```

<!-- grade_id: problem_3_select_data_test, points: 0.5, grade: true, solution: false -->
```python
# NON-EDITABLE CODE CELL FOR TESTING YOUR SOLUTION
print(list(data.columns))
```

Output:
```
['from_x', 'from_y', 'to_x', 'to_y']
```

3: Create (two) empty lists for points called **`orig_points`** and **`dest_points`**


<!-- grade_id: create_lists, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
orig_points = []
dest_points = []
# raise NotImplementedError()
```

Output:
```
NotImplementedError                       Traceback (most recent call last)
<ipython-input-34-e9793ea44b50> in <module>
      2 orig_points = []
      3 dest_points = []
----> 4 raise NotImplementedError()

NotImplementedError:
```

<!-- grade_id: problem_3_create_lists_test, points: 0.5, grade: true, solution: false -->
```python
# NON-EDITABLE CODE CELL FOR TESTING YOUR SOLUTION

# List length should be zero at this point:
print('orig_points length:', len(orig_points))
print('dest_points length:', len(dest_points))
```

Output:
```
orig_points length: 0
dest_points length: 0
```

4: Create shapely points for each origin and destination and add origin points to `orig_points` list and destination points to `dest_points` list.

**HOW?:**

- Create a for-loop and iterate over the rows of your dataframe
- For each row, create Shapely Point -objects based on the coordinate columns (columns `from_x` and `from_y` for the origins and columns `to_x` and `to_y` for the destinations)
- Append the point objects into the **`orig_points`** -list and **`dest_point`** -list.

See Geo-Python Lesson 6 materials [for iterating data frame fows](https://geo-python.github.io/site/notebooks/L6/pandas/advanced-data-processing-with-pandas.html#Iterating-rows-and-using-self-made-functions-in-Pandas) if you do not remember how to do it.

<!-- grade_id: create-points, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE

for i, row in data.iterrows():
    print(i,row['from_x'])
    or_point = Point(row['from_x'], row['from_y'])
    de_point = Point(row['to_x']  , row['to_y'])
    orig_points.append(or_point)
    dest_points.append(de_point)
# raise NotImplementedError()
```

Output:
```
0 24.9704379
1 24.9704379
2 24.9704379
3 24.9704379
4 24.9704379
5 24.9704379
6 24.9704379
7 24.9704379
8 24.9704379
9 24.9704379
10 24.9704379
11 24.9704379
12 24.9704379
13 24.9704379
14 24.9704379
15 24.9704379
16 24.9704379
17 24.9704379
18 24.9704379
19 24.9704379
20 24.9704379
21 24.9704379
22 24.9704379
23 24.9704379
24 24.9704379
25 24.9704379
26 24.9704379
27 24.9704379
28 24.9704379
29 24.9704379
30 24.9704379
31 24.9704379
32 24.9704379
33 24.9704379
34 24.9704379
35 24.9704379
36 24.9704379
37 24.9704379
38 24.9704379
39 24.9704379
40 24.9704379
41 24.9704379
42 24.9704379
43 24.9704379
44 24.9704379
45 24.9704379
46 24.9704379
47 24.9704379
48 24.9704379
49 24.9704379
50 24.9704379
51 24.9704379
52 24.9704379
53 24.9704379
54 24.9704379
55 24.9704379
56 24.9704379
57 24.9704379
58 24.9704379
59 24.9704379
60 24.9704379
61 24.9704379
62 24.9704379
63 24.9704379
64 24.9704379
65 24.9704379
66 24.9704379
67 24.9704379
68 24.9704379
69 24.9704379
70 24.9704379
71 24.9704379
72
```

**NOTE: After you have solved this problem, we recommend that you restart the kernel and run all cells again! Otherwise you might append the same coordinates to the lists many times if you run the cell multiple times.**

<!-- grade_id: problem_3_create_points_test, points: 3.5, grade: true, solution: false -->
```python
# NON-EDITABLE CODE CELL FOR TESTING YOUR SOLUTION

# This test print should print out the first origin and destination coordinates in the two lists:
print("ORIGIN X Y:", orig_points[0].x, orig_points[0].y)
print("DESTINATION X Y:", dest_points[0].x, dest_points[0].y)

#Check that you created a correct amount of points:
assert len(orig_points) == len(data), "Number of origin points must be the same as number of rows in the original file"
assert len(dest_points) == len(data), "Number of destination points must be the same as number of rows in the original file"
```

Output:
```
ORIGIN X Y: 24.9704379 60.3119173
DESTINATION X Y: 24.8560344 60.3999406
```


- Upload the codes and edits to your **own** personal GitHub repository for Exercise-1 in AutoGIS-2018.

## Done!

That's it. Now you are ready to continue for the final Problem 4.

## Problem 4: Creating LineStrings that represent the movements (*5 points*):

This task continuous where we left in Problem 3. 
   
1: Create a list called `lines`


<!-- grade_id: line_list, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
lines = []
# raise NotImplementedError()
```

<!-- grade_id: problem_4_line_list_test, points: 0.5, grade: true, solution: false -->
```python
# NON-EDITABLE CODE CELL FOR TESTING YOUR SOLUTION

# Lines length should be zero at this stage:
print('lines length:', len(lines))
```

Output:
```
lines length: 0
```

2a: Iterate over the origin and destination lists and create a Shapely LineString -object between the origin and destination point

  - Hint - Alternative 1: You can take advantage of `range()` function here that can help accessing the values from two lists at the same time.
     
  - Hint - Alternative 2: You can use `zip()` function to iterate over many lists at the same time. [See hints for this week](https://automating-gis-processes.github.io/2018/lessons/L1/ex-1.html#hints)
  
2b: Add the LineString into the `lines` -list.


<!-- grade_id: create_od_line, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
# points = [Point(45.2, 22.34), Point(100.22, -3.20)] LineString(points)
for i in range(0, len(orig_points)):
    line = LineString([orig_points[i], dest_points[i]])
    lines.append(line)
# raise NotImplementedError()lines.append(Point(point[0],point[1]))

```

**NOTE: After you have solved this problem, we recommend that you restart the kernel and run all cells again! Otherwise you might append the same points to the lists many times if you run the cell multiple times.**

<!-- grade_id: problem_4_create_od_line_test, points: 2, grade: true, solution: false -->
```python
# NON-EDITABLE CODE CELL FOR TESTING YOUR SOLUTION

#Test that the list has correct number of LineStrings
assert len(lines) == len(data), "There should be as many lines as there are rows in the original data"
```

3: Create a variable called **`total_length`**, and store the total (Euclidian) distance of all the origin-destination LineStrings that we just created into that variable.

  - Hint: You might want to iterate over the lines and update the total lenght on each iteration.


<!-- grade_id: euclidian_distance, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
total_length = 0.0
for line in lines:
    total_length = total_length + line.length
# raise NotImplementedError()
```

Output:
```
3148.5707808651528
```

<!-- grade_id: problem_4_euclidian_distance_test, points: 1.5, grade: true, solution: false -->
```python
# NON-EDITABLE CODE CELL FOR TESTING YOUR SOLUTION

# This test print should print the total length of all lines
print("Total length of all lines is", round(total_length, 2))
```

Output:
```
Total length of all lines is 3148.57
```

4: write the previous parts, i.e. the creation of the LineString and calculating the total distance, into dedicated functions:  

- `create_od_lines()`: Takes two lists of Shapely Point -objects as input and returns a list of LineStrings
- `calculate_total_distance()`: Takes a list of LineString geometries as input and returs their total length

You can copy and paste the codes you have written earlier into the functions. Below, you can find a code cell for testing your functions (you should get the same result as earler).

<!-- grade_id: distance_function, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
resulted_lines = []
def create_od_lines(from_list, to_list):
    """
    returns a list of lines using the two points lists in takes as parameters.
    """
    assert type(from_list) == list and type(to_list) == list, 'Input should be a list!'
    for i in range(0, len(from_list)):        
        assert type(from_list[i]) == Point and type(to_list[i]) == Point, 'All list values should be Shapely Point objects!'
        line = LineString([from_list[i], to_list[i]])
        resulted_lines.append(line)
    return resulted_lines

def calculate_total_distance(od_lines):
    total_length = 0.0
    assert type(od_lines) == list, 'Input should be a list!'
    for line in lines:
        assert type(line) == LineString, 'all list values should be a LineString objects'
        total_length = total_length + line.length
    return total_length
# raise NotImplementedError()
```

<!-- grade_id: problem_4_distance_function, points: 1, grade: true, solution: false -->
```python
# NON-EDITABLE CODE CELL FOR TESTING YOUR SOLUTION

# Use the functions
# -----------------

# Create origin-destination lines
od_lines = create_od_lines(orig_points, dest_points)

# Calculate the total distance
tot_dist = calculate_total_distance(od_lines)

print("Total distance", round(tot_dist,2))

```

Output:
```
Total distance 3148.57
```


## All done!

Awesome, now you have successfully practiced how geometries can be created in Python. Next week we will start using them actively.